# Read the Pallet Movers File on Canvas

If you have not read the PDF file that accompanies this assignment found on Blackboard, you will most likely be **lost**.

## Run the Next 2 Code Cells

Run the next two code cells before attempting to complete the tasks. After running both code cells, you will have the following variables in memory and available for use in your tasks below.

- A `list` variable named `hubs` that contains the cities where the hubs are located.
- A `list` variable named `spokes` that contains the cities where the spokes are located.
- A `list` variable named `requests` that contain the requests from potential customers for Pallet Movers to move their freight.

After running the code cell, you should examine these lists to understand their structure, content, etc.

```python
!pip install nose
```

```python
from nose.tools import assert_equal, assert_in, assert_is_instance
from nose.tools import assert_almost_equal, assert_not_equal
```

```python
import pandas as pd
hubs = ['Pittsburgh', 'Newark', 'Richmond', 'Charlotte', 'Atlanta']
spokes = ['Akron', 'Youngstown', 'Allentown', 'Harrisburg', 'Morgantown',
          'White Plains', 'Trenton', 'Philadelphia', 'Elkton', 'Catskill',
          'Charlottesville', 'Lynchburg', 'Newport News', 'Fredericksburg',
          'South Hill', 'Greensboro', 'Fayetteville', 'Columbia',
          'Spartanburg', 'Athens', 'Augusta', 'Chattanooga', 'Birmingham']
requests = pd.read_csv('./data/requests.csv').to_dict(orient='records')
print(f'The variable requests has a type of {type(requests)}')
```

Output:
```
The variable requests has a type of <class 'list'>
```

# 1. Find a Request

Implement the function `find_request` below that takes a request `id` as an argument and returns the entire request as a dictionary. Your function should assume that the variable `requests` is in memory. You can also assume that every request has a unique id. Be sure to read the docstring of the function.

<!-- grade_id: cell-eaf806e613e31fa1, points: 0, grade: false, solution: true -->
```python
def find_request(id):
    """
    A function that returns a shipping request from the list `requests`
    
    Parameters
    ----------
    id: the string id of a request to find
    
    Returns
    -------
    request: the request as a dictionary. If the request is not found return None
    """
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-b92446e5a2c5fb82, points: 10, grade: true, solution: false -->
```python
# Testing your function
r1 = find_request('2023_01_0005')
assert_equal(type(r1), dict, msg='Your function did not return a dictionary as required')
assert_equal(r1['destination_city'], 'Allentown', msg='Your function did not return the right output')
```

<!-- grade_id: cell-0333beffd2361ec9, points: 10, grade: true, solution: false -->
```python
# Testing your function
r2 = find_request('junk')
assert_equal(r2, None, msg='Your function did not return None when the id was not found')
```

# 2. Find Requests Based on Origin City

You want to find all the requests based on a specific origin city found in the request. Implement the function `get_requests_by_origin_city` below. The function should return a list of dictionaries as indicated in the docstring. Your search must be case insensitive. 

<!-- grade_id: cell-57d13dcbe9bacec9, points: 0, grade: false, solution: true -->
```python
def get_requests_by_origin_city(origin_city):
    """
    A function that takes in the name of the origin city and returns 
    a list of all the requests for that particular origin city. The
    search is case insensitive.
    
    Parameters
    ----------
    origin_city: a string representing the origin city to search for
    
    Returns
    -------
    results: a list of requests, where each request is a dictionary.
             If no requests exist for the origin city, then return
            an empty list.
    """
    ### YOUR CODE HERE
    
```

<!-- grade_id: cell-8b6d05c594987254, points: 10, grade: true, solution: false -->
```python
# Testing your function
richmond = get_requests_by_origin_city('ricHMond')
assert_equal(type(richmond), list, msg='Your function did not return a list as required')
assert_equal(len(richmond), 13, msg='Your function did not return the correct number of requests')
assert_equal(richmond[5]['request_id'], '2023_01_0113', msg='Your function did not return the right output')
```

<!-- grade_id: cell-a7f456f9daea3eb2, points: 10, grade: true, solution: false -->
```python
# Testing your function
bad_city = get_requests_by_origin_city('junk')
assert_equal(type(bad_city), list, msg='Your function did not return a list the origin city was not found')
assert_equal(len(bad_city), 0, msg='Your function did not return an empty list the origin city was not found')
```

# 3. Finding Automatically Quotable Requests

Currently, Pallet Movers manually reviews every request to generate a quote for the customer. A quote generally involves a price and estimated pickup and delivery dates based on the information that the customer provided in their request. Pallet Movers wants to streamline some of this manually intensive quoting process. Their IT department has devised a way to generate automatic quotes for a subset of requests. To use the automated quoting tool, both the origin and destination have to be part of their hub-and-spoke system. If either the origin or destination city is not in their network, the tool explodes and catches the home office on fire. Obviously, you want to avoid this situation because of your aversion to fire. 

You have been tasked to write a function that sifts through the requests, separating them into four lists. The first list should contain the requests that can be automatically quoted. The second list should contain the requests that only have an origin city **outside** of Pallet Movers network. The third list should contain the requests that only have a destination city **outside** of their network. The fourth list should contain the requests that have **both** the origin and destination cities outside of the network.

Name your function `sort_requests`. It should have three arguments: 

1. `the_requests` which the big list of request to sort through
2. `the_hubs` which is a list of the hub cities
3. `the_spokes` which is a list of the spoke cities

As noted above, the function should return all four lists in the order specified. You must also make sure that all of your sorting and searching is case insensitive to protect against a customer inadvertantly not using proper case when making the request.

<!-- grade_id: cell-590b55a66c4e910e, points: 0, grade: false, solution: true -->
```python
# Create your user-defined function here
### YOUR CODE HERE

```

```python
# Test your function
good, bad_o, bad_d, both_bad = sort_requests(requests, hubs, spokes)
print(f'The number of quotable requests is {len(good)}')
print(f'The number of bad origin requests is {len(bad_origin)}')
print(f'The number of bad destination requests is {len(bad_dest)}')
print(f'The number of both bad origin and destination requests is {len(both_bad)}')
```

<!-- grade_id: cell-3d2ff214ba63763f, points: 7, grade: true, solution: false -->
```python
# Test your function
good, bad_o, bad_d, both_bad = sort_requests(requests, hubs, spokes)

assert_equal(type(good), list, msg='Your function did not return a list for the first element of the tuple as required')
assert_equal(type(good[-1]), dict, msg='The good list does not contain dictionaries')
assert_equal(len(good), 306, msg='Your function did not return the correct number of good requests')
```

<!-- grade_id: cell-14bc4ff37e7568aa, points: 7, grade: true, solution: false -->
```python
# Test your function
good, bad_o, bad_d, both_bad = sort_requests(requests, hubs, spokes)

assert_equal(type(bad_o), list, msg='Your function did not return a list for the second element of the tuple as required')
assert_equal(type(bad_o[0]), dict, msg='The bad origin list does not contain dictionaries')
assert_equal(len(bad_o), 5, msg='Your function did not return the correct number of bad origin requests')
```

<!-- grade_id: cell-374e2b8befd481a8, points: 7, grade: true, solution: false -->
```python
# Test your function
good, bad_o, bad_d, both_bad = sort_requests(requests, hubs, spokes)

assert_equal(type(bad_d), list, msg='Your function did not return a list for the third element of the tuple as required')
assert_equal(type(bad_d[0]), dict, msg='The bad destination list does not contain dictionaries')
assert_equal(len(bad_d), 4, msg='Your function did not return the correct number of bad destination requests')
```

<!-- grade_id: cell-91d9c7e7c5d83d16, points: 7, grade: true, solution: false -->
```python
# Test your function
good, bad_o, bad_d, both_bad = sort_requests(requests, hubs, spokes)

assert_equal(type(both_bad), list, msg='Your function did not return a list for the fourth element of the tuple as required')
assert_equal(type(both_bad[0]), dict, msg='The both bad list does not contain dictionaries')
assert_equal(len(both_bad), 2, msg='Your function did not return the correct number of both bad requests')
```

<!-- grade_id: cell-0f09fc1bb0a9d2a3, points: 7, grade: true, solution: false -->
```python
# Test your function
good, bad_o, bad_d, both_bad = sort_requests(requests, hubs, spokes)

assert_equal(type(bad_d[1]), dict, msg='The bad destination list does not contain dictionaries')
assert_equal(bad_d[1]['destination_city'], 'Cartersville', msg='Your function did not return the correct name of the of bad destination city')
```

**&copy; 2022 - Present: Matthew D. Dean, Ph.D.   
Clinical Associate Professor of Business Analytics at William \& Mary.**